<a href="https://colab.research.google.com/github/esther119/Algorithm-and-Data-Structure/blob/master/CS110_PCW_6_Asymptotic_notations%3B_maximum_subarray_problem%E2%80%94part_I.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **37 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Esther"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Asymptotic notations; maximum-subarray problem

## Question 1 (Exercise 3.1-3 of Cormen, et al. ) [time estimate: 5 minutes]
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is true but meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
O is to measure the worst case scenario, therefore, at least reaching the worst case scenario doesn't make sense. 

## Question 2 (Exercise 3.1-4 of Cormen, et al. ) [time estimate: 5 minutes]

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$? Why?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->


*   Yes. $2^{n+1}$ can be expressed as $2*2^n$. Since the coefficient 
*   2 is a constant, it is same as O($2^n$). 
No.$2^{2n}$ = $(2^n)* (2^n)$ which is two times larger than $2^n$. 



## Question 3 [time estimate: 20 minutes]
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
     
    """
    marker_sum = -float('inf')
    start_point = 0 
    end_point = 0 
    for i in range(len(A)):
      for j in range(i+1, len(A)): 
        new_sum = sum(A[i:j])
        if marker_sum < new_sum: #if the new array has a sum greater than our marker, we update it 
          marker_sum = new_sum
          start_point = i 
          end_point = j-1

    return start_point, end_point, marker_sum
        

A = [-2, -5, 6, -2, -3, 1, 5, -6]
bruteforce_max_subarray(A)

```

Output:
```
(2, 6, 7)
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

## Divide and conquer approach practice
max sum occurs in three scenarios  
1. entirely at the left sublist
2. entirely at the right sublist
3. crossing the middle point 


```python
def max_crossing_subarray(A, low, mid, high):
  
  left_sum = -float('inf') #start from the middle go toward the left. Find the max at the left hand side
  sum = 0 
  for i in range(mid,low-1,-1):
    sum = sum +A[i]
    if sum > left_sum: 
      left_sum = sum
      left_index = i 
  right_sum = -float('inf') #start from the middle go toward the right. Find the max at the right eft hand side
  sum = 0 
  for j in range(mid+1,high+1):
    sum = sum +A[j]
    if sum > right_sum: 
      right_sum = sum
      right_index = j

  return left_index, right_index, left_sum + right_sum

A = [-2,1,-1,2,-5]
max_crossing_subarray(A,0,2,4)
B = [-2,1,3,4]
max_crossing_subarray(B,0,1,3)
```

Output:
```
(1, 3, 8)
```

```python
def max_subarray(A, low, high):
  if high == low: #base case
    return low, high, A[low]
  else: 
    mid = (low+high)//2
    left_low,left_high, max_left = max_subarray(A, low, mid)
    right_low, right_high,max_right = max_subarray(A, mid+1, high)
    cross_low, cross_high ,max_crossing = max_crossing_subarray(A, low, mid, high) 
    print(max_left, max_right, max_crossing)
    #compare the sum for the left, right hand side of the list and the crossing 
    if max_left >= max_right and max_left > max_crossing:
      return left_low,left_high, max_left
    elif max_right > max_left and max_right > max_crossing:
      return right_low, right_high,max_right
    else: 
      return cross_low, cross_high ,max_crossing


A = [-2,1,-1,2,-5]
max_subarray(A,0,4)

#From the output, we see that we will check the left, right, and cross side, at every level. And pick the largest among the three of them.
```

Output:
```
-2 1 -1
1 -1 0
2 -5 -3
1 2 2
(1, 3, 2)
```

```python
C = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
max_subarray(C,0,len(C)-1)
```

Output:
```
13 -3 10
-25 20 -5
13 20 5
-3 -16 -19
-23 18 -5
-3 18 -21
20 18 17
20 -7 13
12 -5 7
20 12 25
-22 15 -7
-4 7 3
15 7 18
25 18 16
20 25 43
(7, 10, 43)
```

## Question 4 [time estimate: 2 minutes]
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
bruteforce_max_subarray(A)
```

Output:
```
(7, 10, 43)
```

## Question 5. Asymptotic notation [time estimate: 5 minutes]
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Running time 
--- | --- 
Insertion sort | $O(n^2)$ , $\Theta(n)$, $\Omega(n^2)$
Selection sort |  $O(n^2)$ , $\Theta(n^2)$, $\Omega(n^2)$
Bubble sort |  $O(n^2)$ , $\Theta(n)$, $\Omega(n^2)$
Finding maximum subarray(brute force) |  $O(n^2)$ , $\Theta(n)$, $\Omega(n^2)$


*Bubble sort: if we only count swap instead of comparisons 

## [Optional] Question 6 [time estimate: 15 minutes] 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """
    Implements brute-force minimum subarray finding.
    
    """
    marker_sum = +float('inf')
    start_point = 0 
    end_point = 0 
    for i in range(len(A)):
      for j in range(i+1, len(A)): 
        new_sum = sum(A[i:j])
        if marker_sum > new_sum: #if the new array has a sum greater than our marker, we update it 
          marker_sum = new_sum
          start_point = i 
          end_point = j-1

    return start_point, end_point, marker_sum
        

A = [-2, -5, 6, -2, -3, 1, 5, -6]
bruteforce_min_subarray(A)
```

Output:
```
(0, 1, -7)
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```

```python

```